<?xml version="1.0" encoding="UTF-8"?>
<quiz>
<question type="stack">
<name>
<text>Pizzas ht</text>
</name>
<questiontext format="html">
<text><![CDATA[<p>A quality-control officer at Eagle Boys measured the diameter (in inches) of a random sample of \(n={@n@}\) pizzas advertised as having diameter \(12\) inches. The sample mean was \(\bar{x}={@xbar@}\) inches, with sample standard deviation \(s={@s@}\) inches.</p>

<p>We wish to test, at the \(\alpha=0.05\) significance level, whether the true mean diameter differs from \(12\) inches.</p>

<h4>1. Standard error</h4>
<p>Compute the standard error of the sample mean (give your answer correct to at least 3 significant figures).</p>
[[input:ans1]] [[validation:ans1]] [[feedback:prt_ans1]]

<h4>2. Hypotheses</h4>
<p>State the null and alternative hypotheses, in terms of the population mean \(\mu\).</p>
<p>\(H_0:\) [[input:ans2a]] [[validation:ans2a]] [[feedback:prt_ans2a]]</p>
<p>\(H_1:\) [[input:ans2b]] [[validation:ans2b]] [[feedback:prt_ans2b]]</p>

<h4>3. Test statistic</h4>
<p>Using your own answer to Part 1 as the standard error, compute the \(t\)-score for this sample (give your answer to 2 decimal places).</p>
[[input:ans3]] [[validation:ans3]] [[feedback:prt_ans3]]

<h4>4. P-value estimate</h4>
<p>Using the 68-95-99.7 rule and your own \(t\)-score from Part 3, which of the following best describes the two-sided P-value?</p>
[[input:ans4]] [[validation:ans4]] [[feedback:prt_ans4]]

<h4>5. Conclusion</h4>
<p>Based on your own P-value estimate from Part 4, and using \(\alpha=0.05\), what conclusion should be drawn about \(H_0\)?</p>
[[input:ans5]] [[validation:ans5]] [[feedback:prt_ans5]]

<h4>6. Validity check</h4>
<p>Regardless of your previous answers: is the sample size \(n={@n@}\) large enough for the Central Limit Theorem to justify this t-test, using the textbook's rule of thumb that \(n\) should be much larger than 25?</p>
[[input:ans6]] [[validation:ans6]] [[feedback:prt_ans6]]]]></text>
</questiontext>
<questionvariables>
<text><![CDATA[w: rand(100);
bucket_target: if w<50 then 1 elseif w<70 then 2 elseif w<85 then 3 else 4;

t_abs: if bucket_target=1 then 3.2+rand(9)/10
       elseif bucket_target=2 then 2.2+rand(7)/10
       elseif bucket_target=3 then 1.2+rand(7)/10
       else 0.3+rand(6)/10;

s: 0.15+rand(21)/100;
n: rand(35)+15;

xbar_exact: 12 - t_abs*s/sqrt(n);
xbar: float(round(1000*xbar_exact)/1000);

se_true: s/sqrt(n);
t_true: (xbar-12)/se_true;

bucket_true: if abs(t_true)>3 then 1 elseif abs(t_true)>2 then 2 elseif abs(t_true)>1 then 3 else 4;

decision_true: if bucket_true=1 or bucket_true=2 then 1 else 2;

validity_true: if n>25 then 1 else 2;

h0_ta: mu=12;
h1_ta: mu#12;

ta4: [[1, is(bucket_true=1), "\\( P \\lt 0.003 \\)"],
      [2, is(bucket_true=2), "\\( 0.003 \\lt P \\lt 0.05 \\)"],
      [3, is(bucket_true=3), "\\( 0.05 \\lt P \\lt 0.32 \\)"],
      [4, is(bucket_true=4), "\\( P \\gt 0.32 \\)"]];
ta4: random_permutation(ta4);

ta5: [[1, is(decision_true=1), "Reject \\(H_0\\)"],
      [2, is(decision_true=2), "Fail to reject \\(H_0\\)"]];

ta6: [[1, is(validity_true=1), "Yes"],
      [2, is(validity_true=2), "No"]];]]></text>
</questionvariables>
<generalfeedback format="html">
<text><![CDATA[<p>The standard error of the mean is</p>
\[ SE = \frac{s}{\sqrt{n}} = \frac{ {@s@} }{\sqrt{ {@n@} }} = {@se_true@}. \]

<p>The hypotheses are \(H_0:\mu=12\) versus \(H_1:\mu\neq12\).</p>

<p>The test statistic is</p>
\[ t = \frac{\bar{x}-\mu_0}{SE} = \frac{ {@xbar@} - 12}{ {@se_true@} } = {@t_true@}. \]

[[ if test="bucket_true=1" ]]
<p>Since \(|t| > 3\), by the 68-95-99.7 rule the two-sided P-value satisfies \(P<0.003\).</p>
[[ elif test="bucket_true=2" ]]
<p>Since \(2<|t|<3\), by the 68-95-99.7 rule the two-sided P-value satisfies \(0.003<P<0.05\).</p>
[[ elif test="bucket_true=3" ]]
<p>Since \(1<|t|<2\), by the 68-95-99.7 rule the two-sided P-value satisfies \(0.05<P<0.32\).</p>
[[ else ]]
<p>Since \(|t|<1\), by the 68-95-99.7 rule the two-sided P-value satisfies \(P>0.32\).</p>
[[/ if ]]

[[ if test="decision_true=1" ]]
<p>Since \(P<0.05\), we reject \(H_0\): there is evidence the mean diameter differs from 12 inches.</p>
[[ else ]]
<p>Since \(P>0.05\), we fail to reject \(H_0\): there is not enough evidence that the mean diameter differs from 12 inches.</p>
[[/ if ]]

[[ if test="validity_true=1" ]]
<p>Since \(n={@n@}\) is much larger than 25, the sample size is large enough for the Central Limit Theorem to justify this t-test.</p>
[[ else ]]
<p>Since \(n={@n@}\) is not much larger than 25, the sample size is not large enough to safely justify this t-test using the Central Limit Theorem.</p>
[[/ if ]]]]></text>
</generalfeedback>
<questionnote format="moodle_auto_format">
<text>\(\bar{x}={@xbar@}\), \(s={@s@}\), \(n={@n@}\), \(t={@t_true@}\), bucket={@bucket_true@}, decision={@decision_true@}, validity={@validity_true@}</text>
</questionnote>
<input>
<name>ans1</name>
<type>numerical</type>
<tans>se_true</tans>
<forbidfloat>0</forbidfloat>
</input>
<input>
<name>ans2a</name>
<type>algebraic</type>
<tans>h0_ta</tans>
</input>
<input>
<name>ans2b</name>
<type>algebraic</type>
<tans>h1_ta</tans>
</input>
<input>
<name>ans3</name>
<type>numerical</type>
<tans>t_true</tans>
<forbidfloat>0</forbidfloat>
</input>
<input>
<name>ans4</name>
<type>radio</type>
<tans>ta4</tans>
<showvalidation>0</showvalidation>
</input>
<input>
<name>ans5</name>
<type>radio</type>
<tans>ta5</tans>
<showvalidation>0</showvalidation>
</input>
<input>
<name>ans6</name>
<type>radio</type>
<tans>ta6</tans>
<showvalidation>0</showvalidation>
</input>
<prt>
  <name>prt_ans1</name>
  <node>
    <name>0</name>
    <answertest>AlgEquiv</answertest>
    <sans>is(abs(ans1-se_true)&lt;=0.01*se_true)</sans>
    <tans>true</tans>
    <trueanswernote>prt_ans1-1-T</trueanswernote>
    <falsenextnode>1</falsenextnode>
    <falseanswernote>prt_ans1-1-F</falseanswernote>
  </node>
  <node>
    <name>1</name>
    <answertest>AlgEquiv</answertest>
    <sans>is(abs(ans1-(s/n))&lt;=0.01*(s/n))</sans>
    <tans>true</tans>
    <truescore>0</truescore>
    <trueanswernote>prt_ans1-2-T</trueanswernote>
    <truefeedback format="html">
      <text><![CDATA[<p>It looks like you divided \(s\) by \(n\) instead of \(\sqrt{n}\). Recall \(SE=s/\sqrt{n}\).</p>]]></text>
    </truefeedback>
    <falseanswernote>prt_ans1-2-F</falseanswernote>
    <falsefeedback format="html">
      <text><![CDATA[<p>Not correct — check the formula \(SE=s/\sqrt{n}\).</p>]]></text>
    </falsefeedback>
  </node>
</prt>

<prt>
  <name>prt_ans2a</name>
  <node>
    <name>0</name>
    <answertest>AlgEquiv</answertest>
    <sans>ans2a</sans>
    <tans>h0_ta</tans>
    <trueanswernote>prt_ans2a-1-T</trueanswernote>
    <falseanswernote>prt_ans2a-1-F</falseanswernote>
  </node>
</prt>

<prt>
  <name>prt_ans2b</name>
  <node>
    <name>0</name>
    <answertest>AlgEquiv</answertest>
    <sans>ans2b</sans>
    <tans>h1_ta</tans>
    <trueanswernote>prt_ans2b-1-T</trueanswernote>
    <falseanswernote>prt_ans2b-1-F</falseanswernote>
  </node>
</prt>

<prt>
  <name>prt_ans3</name>
  <node>
    <name>0</name>
    <answertest>AlgEquiv</answertest>
    <sans>is(ans1=0)</sans>
    <tans>true</tans>
    <truescore>0</truescore>
    <trueanswernote>prt_ans3-1-T</trueanswernote>
    <truefeedback format="html">
      <text><![CDATA[<p>Your standard error is zero, so a \(t\)-score cannot be computed from it.</p>]]></text>
    </truefeedback>
    <falsenextnode>1</falsenextnode>
    <falseanswernote>prt_ans3-1-F</falseanswernote>
  </node>
  <node>
    <name>1</name>
    <answertest>AlgEquiv</answertest>
    <sans>is(abs(ans3-(xbar-12)/ans1)&lt;=0.02)</sans>
    <tans>true</tans>
    <trueanswernote>prt_ans3-2-T</trueanswernote>
    <truefeedback format="html">
      <text><![CDATA[<p>Correct — well done.</p>]]></text>
    </truefeedback>
    <falsenextnode>2</falsenextnode>
    <falseanswernote>prt_ans3-2-F</falseanswernote>
  </node>
  <node>
    <name>2</name>
    <answertest>AlgEquiv</answertest>
    <sans>is(abs(ans3-(12-xbar)/ans1)&lt;=0.02)</sans>
    <tans>true</tans>
    <truescore>0</truescore>
    <trueanswernote>prt_ans3-3-T</trueanswernote>
    <truefeedback format="html">
      <text><![CDATA[<p>It looks like you flipped the sign — use \(t=(\bar{x}-\mu_0)/SE\), not \((\mu_0-\bar{x})/SE\).</p>]]></text>
    </truefeedback>
    <falseanswernote>prt_ans3-3-F</falseanswernote>
    <falsefeedback format="html">
      <text><![CDATA[<p>Not correct — recompute \(t=(\bar{x}-\mu_0)/SE\) using your own answer to Part 1 as \(SE\).</p>]]></text>
    </falsefeedback>
  </node>
</prt>

<prt>
  <name>prt_ans4</name>
  <node>
    <name>0</name>
    <answertest>AlgEquiv</answertest>
    <sans>is(ans4=(if abs(ans3)>3 then 1 elseif abs(ans3)>2 then 2 elseif abs(ans3)>1 then 3 else 4))</sans>
    <tans>true</tans>
    <trueanswernote>prt_ans4-1-T</trueanswernote>
    <truefeedback format="html">
      <text><![CDATA[<p>Correct bucket, consistent with your own \(t\)-score.</p>]]></text>
    </truefeedback>
    <falseanswernote>prt_ans4-1-F</falseanswernote>
    <falsefeedback format="html">
      <text><![CDATA[<p>Not correct — re-apply the 68-95-99.7 rule to your own \(t\)-score from Part 3.</p>]]></text>
    </falsefeedback>
  </node>
</prt>

<prt>
  <name>prt_ans5</name>
  <node>
    <name>0</name>
    <answertest>AlgEquiv</answertest>
    <sans>is(ans5=(if ans4=1 or ans4=2 then 1 else 2))</sans>
    <tans>true</tans>
    <trueanswernote>prt_ans5-1-T</trueanswernote>
    <truefeedback format="html">
      <text><![CDATA[<p>Correct conclusion, consistent with your own P-value bucket.</p>]]></text>
    </truefeedback>
    <falseanswernote>prt_ans5-1-F</falseanswernote>
    <falsefeedback format="html">
      <text><![CDATA[<p>Not correct — your conclusion should match the P-value bucket you chose in Part 4: reject if \(P<0.05\), otherwise fail to reject.</p>]]></text>
    </falsefeedback>
  </node>
</prt>

<prt>
  <name>prt_ans6</name>
  <node>
    <name>0</name>
    <answertest>AlgEquiv</answertest>
    <sans>is(ans6=validity_true)</sans>
    <tans>true</tans>
    <trueanswernote>prt_ans6-1-T</trueanswernote>
    <truefeedback format="html">
      <text><![CDATA[<p>Correct.</p>]]></text>
    </truefeedback>
    <falseanswernote>prt_ans6-1-F</falseanswernote>
    <falsefeedback format="html">
      <text><![CDATA[<p>Not correct — compare \(n\) to the textbook's rule of thumb of \(n\) much larger than 25.</p>]]></text>
    </falsefeedback>
  </node>
</prt>
<qtest>
  <testcase>1</testcase>
  <description>All correct (teacher's answer)</description>
  <testinput>
    <name>ans1</name>
    <value>se_true</value>
  </testinput>
  <testinput>
    <name>ans2a</name>
    <value>h0_ta</value>
  </testinput>
  <testinput>
    <name>ans2b</name>
    <value>h1_ta</value>
  </testinput>
  <testinput>
    <name>ans3</name>
    <value>t_true</value>
  </testinput>
  <testinput>
    <name>ans4</name>
    <value>bucket_true</value>
  </testinput>
  <testinput>
    <name>ans5</name>
    <value>decision_true</value>
  </testinput>
  <testinput>
    <name>ans6</name>
    <value>validity_true</value>
  </testinput>
  <expected>
    <name>prt_ans1</name>
    <expectedscore>1.0000000</expectedscore>
    <expectedpenalty>0.0000000</expectedpenalty>
    <expectedanswernote>prt_ans1-1-T</expectedanswernote>
  </expected>
  <expected>
    <name>prt_ans2a</name>
    <expectedscore>1.0000000</expectedscore>
    <expectedpenalty>0.0000000</expectedpenalty>
    <expectedanswernote>prt_ans2a-1-T</expectedanswernote>
  </expected>
  <expected>
    <name>prt_ans2b</name>
    <expectedscore>1.0000000</expectedscore>
    <expectedpenalty>0.0000000</expectedpenalty>
    <expectedanswernote>prt_ans2b-1-T</expectedanswernote>
  </expected>
  <expected>
    <name>prt_ans3</name>
    <expectedscore>1.0000000</expectedscore>
    <expectedpenalty>0.0000000</expectedpenalty>
    <expectedanswernote>prt_ans3-2-T</expectedanswernote>
  </expected>
  <expected>
    <name>prt_ans4</name>
    <expectedscore>1.0000000</expectedscore>
    <expectedpenalty>0.0000000</expectedpenalty>
    <expectedanswernote>prt_ans4-1-T</expectedanswernote>
  </expected>
  <expected>
    <name>prt_ans5</name>
    <expectedscore>1.0000000</expectedscore>
    <expectedpenalty>0.0000000</expectedpenalty>
    <expectedanswernote>prt_ans5-1-T</expectedanswernote>
  </expected>
  <expected>
    <name>prt_ans6</name>
    <expectedscore>1.0000000</expectedscore>
    <expectedpenalty>0.0000000</expectedpenalty>
    <expectedanswernote>prt_ans6-1-T</expectedanswernote>
  </expected>
</qtest>

<qtest>
  <testcase>2</testcase>
  <description>SE computed as s/n (common mistake), but t-score correctly follows through from that wrong SE</description>
  <testinput>
    <name>ans1</name>
    <value>s/n</value>
  </testinput>
  <testinput>
    <name>ans2a</name>
    <value>h0_ta</value>
  </testinput>
  <testinput>
    <name>ans2b</name>
    <value>h1_ta</value>
  </testinput>
  <testinput>
    <name>ans3</name>
    <value>(xbar-12)/(s/n)</value>
  </testinput>
  <testinput>
    <name>ans4</name>
    <value>bucket_true</value>
  </testinput>
  <testinput>
    <name>ans5</name>
    <value>decision_true</value>
  </testinput>
  <testinput>
    <name>ans6</name>
    <value>validity_true</value>
  </testinput>
  <expected>
    <name>prt_ans1</name>
    <expectedscore>0.0000000</expectedscore>
    <expectedpenalty>0.1000000</expectedpenalty>
    <expectedanswernote>prt_ans1-2-T</expectedanswernote>
  </expected>
  <expected>
    <name>prt_ans3</name>
    <expectedscore>1.0000000</expectedscore>
    <expectedpenalty>0.0000000</expectedpenalty>
    <expectedanswernote>prt_ans3-2-T</expectedanswernote>
  </expected>
</qtest>

<qtest>
  <testcase>3</testcase>
  <description>Correct t-score but wrong bucket chosen in Part 4</description>
  <testinput>
    <name>ans1</name>
    <value>se_true</value>
  </testinput>
  <testinput>
    <name>ans3</name>
    <value>t_true</value>
  </testinput>
  <testinput>
    <name>ans4</name>
    <value>mod(bucket_true,4)+1</value>
  </testinput>
  <expected>
    <name>prt_ans3</name>
    <expectedscore>1.0000000</expectedscore>
    <expectedpenalty>0.0000000</expectedpenalty>
    <expectedanswernote>prt_ans3-2-T</expectedanswernote>
  </expected>
  <expected>
    <name>prt_ans4</name>
    <expectedscore>0.0000000</expectedscore>
    <expectedpenalty>0.1000000</expectedpenalty>
    <expectedanswernote>prt_ans4-1-F</expectedanswernote>
  </expected>
</qtest>

<qtest>
  <testcase>4</testcase>
  <description>Wrong bucket in Part 4, but Part 5 decision follows through from that wrong bucket (not the true one)</description>
  <testinput>
    <name>ans1</name>
    <value>se_true</value>
  </testinput>
  <testinput>
    <name>ans3</name>
    <value>t_true</value>
  </testinput>
  <testinput>
    <name>ans4</name>
    <value>mod(bucket_true,4)+1</value>
  </testinput>
  <testinput>
    <name>ans5</name>
    <value>if (mod(bucket_true,4)+1)=1 or (mod(bucket_true,4)+1)=2 then 1 else 2</value>
  </testinput>
  <expected>
    <name>prt_ans4</name>
    <expectedscore>0.0000000</expectedscore>
    <expectedpenalty>0.1000000</expectedpenalty>
    <expectedanswernote>prt_ans4-1-F</expectedanswernote>
  </expected>
  <expected>
    <name>prt_ans5</name>
    <expectedscore>1.0000000</expectedscore>
    <expectedpenalty>0.0000000</expectedpenalty>
    <expectedanswernote>prt_ans5-1-T</expectedanswernote>
  </expected>
</qtest>

<qtest>
  <testcase>5</testcase>
  <description>Validity check wrong, independent of other parts</description>
  <testinput>
    <name>ans6</name>
    <value>if validity_true=1 then 2 else 1</value>
  </testinput>
  <expected>
    <name>prt_ans6</name>
    <expectedscore>0.0000000</expectedscore>
    <expectedpenalty>0.1000000</expectedpenalty>
    <expectedanswernote>prt_ans6-1-F</expectedanswernote>
  </expected>
</qtest>

<qtest>
  <testcase>6</testcase>
  <description>t-score boundary just above 3 -> bucket 1</description>
  <testinput>
    <name>ans3</name>
    <value>3.01</value>
  </testinput>
  <testinput>
    <name>ans4</name>
    <value>1</value>
  </testinput>
  <expected>
    <name>prt_ans4</name>
    <expectedscore>1.0000000</expectedscore>
    <expectedpenalty>0.0000000</expectedpenalty>
    <expectedanswernote>prt_ans4-1-T</expectedanswernote>
  </expected>
</qtest>

<qtest>
  <testcase>7</testcase>
  <description>t-score boundary just below 3 -> bucket 2</description>
  <testinput>
    <name>ans3</name>
    <value>2.99</value>
  </testinput>
  <testinput>
    <name>ans4</name>
    <value>2</value>
  </testinput>
  <expected>
    <name>prt_ans4</name>
    <expectedscore>1.0000000</expectedscore>
    <expectedpenalty>0.0000000</expectedpenalty>
    <expectedanswernote>prt_ans4-1-T</expectedanswernote>
  </expected>
</qtest>

<qtest>
  <testcase>8</testcase>
  <description>t-score boundary just below 2 -> bucket 3</description>
  <testinput>
    <name>ans3</name>
    <value>1.99</value>
  </testinput>
  <testinput>
    <name>ans4</name>
    <value>3</value>
  </testinput>
  <expected>
    <name>prt_ans4</name>
    <expectedscore>1.0000000</expectedscore>
    <expectedpenalty>0.0000000</expectedpenalty>
    <expectedanswernote>prt_ans4-1-T</expectedanswernote>
  </expected>
</qtest>

<qtest>
  <testcase>9</testcase>
  <description>t-score boundary just below 1 -> bucket 4</description>
  <testinput>
    <name>ans3</name>
    <value>0.99</value>
  </testinput>
  <testinput>
    <name>ans4</name>
    <value>4</value>
  </testinput>
  <expected>
    <name>prt_ans4</name>
    <expectedscore>1.0000000</expectedscore>
    <expectedpenalty>0.0000000</expectedpenalty>
    <expectedanswernote>prt_ans4-1-T</expectedanswernote>
  </expected>
</qtest>

<qtest>
  <testcase>10</testcase>
  <description>Blank/zero SE guard on Part 3 — division-by-zero adversarial case</description>
  <testinput>
    <name>ans1</name>
    <value>0</value>
  </testinput>
  <testinput>
    <name>ans3</name>
    <value>0</value>
  </testinput>
  <expected>
    <name>prt_ans1</name>
    <expectedscore>0.0000000</expectedscore>
    <expectedpenalty>0.1000000</expectedpenalty>
    <expectedanswernote>prt_ans1-2-F</expectedanswernote>
  </expected>
  <expected>
    <name>prt_ans3</name>
    <expectedscore>0.0000000</expectedscore>
    <expectedpenalty>0.1000000</expectedpenalty>
    <expectedanswernote>prt_ans3-1-T</expectedanswernote>
  </expected>
</qtest>

<qtest>
  <testcase>11</testcase>
  <description>Sign-flip mistake on t-score, using correct own SE</description>
  <testinput>
    <name>ans1</name>
    <value>se_true</value>
  </testinput>
  <testinput>
    <name>ans3</name>
    <value>(12-xbar)/se_true</value>
  </testinput>
  <expected>
    <name>prt_ans1</name>
    <expectedscore>1.0000000</expectedscore>
    <expectedpenalty>0.0000000</expectedpenalty>
    <expectedanswernote>prt_ans1-1-T</expectedanswernote>
  </expected>
  <expected>
    <name>prt_ans3</name>
    <expectedscore>0.0000000</expectedscore>
    <expectedpenalty>0.1000000</expectedpenalty>
    <expectedanswernote>prt_ans3-3-T</expectedanswernote>
  </expected>
</qtest>
</question>
</quiz>