<?xml version="1.0" encoding="UTF-8"?>
<quiz>
<question type="stack">
<name>
<text>Ethiopian farmers</text>
</name>
<questiontext format="html">
<text><![CDATA[<p>A study compares the adoption of a new technology between small firms and large firms. The observed data are:</p>

<table border="1" cellpadding="6" style="border-collapse:collapse; text-align:center;">
<tr><th></th><th>Adopters</th><th>Non-adopters</th><th>Total</th></tr>
<tr><th>Small firms</th><td>{@A@}</td><td>{@B@}</td><td>{@A+B@}</td></tr>
<tr><th>Large firms</th><td>{@C@}</td><td>{@D@}</td><td>{@C+D@}</td></tr>
</table>

<h4>Thread 1: Difference in proportions</h4>

<p>(a) Which of the following correctly defines the quantity we compare in a two-proportion test here?</p>
[[input:ans1a]] [[validation:ans1a]] [[feedback:prt_ans1a]]

<p>(b) Enter the sample proportion of non-adopters in each group, as a list \([\hat p_{\text{small}}, \hat p_{\text{large}}]\), to at least 3 decimal places.</p>
[[input:ans1b]] [[validation:ans1b]] [[feedback:prt_ans1b]]

<p>(c) Using your own values from part (b), compute the pooled two-proportion \(z\)-statistic for testing \(H_0: p_{\text{small}}=p_{\text{large}}\).</p>
[[input:ans1c]] [[validation:ans1c]] [[feedback:prt_ans1c]]

<p>(d) Based on your \(z\)-statistic from part (c), and a significance level where the critical value is \(\pm 1.96\), what is your conclusion?</p>
[[input:ans1d]] [[validation:ans1d]] [[feedback:prt_ans1d]]

<p>(e) Using your values from part (b), give a 95% confidence interval for \(p_{\text{small}}-p_{\text{large}}\) (unpooled standard error), as a list \([\text{lower},\text{upper}]\).</p>
[[input:ans1e]] [[validation:ans1e]] [[feedback:prt_ans1e]]

<h4>Thread 2: Odds ratio</h4>

<p>(a) Which of the following correctly defines the odds ratio being compared here?</p>
[[input:ans2a]] [[validation:ans2a]] [[feedback:prt_ans2a]]

<p>(b) Enter the odds of non-adoption (non-adopters ÷ adopters) in each group, as a list \([\text{odds}_{\text{small}}, \text{odds}_{\text{large}}]\), to at least 3 decimal places.</p>
[[input:ans2b]] [[validation:ans2b]] [[feedback:prt_ans2b]]

<p>(c) Using your own values from part (b), compute the chi-square test statistic for independence in this table.</p>
[[input:ans2c]] [[validation:ans2c]] [[feedback:prt_ans2c]]

<p>(d) Based on your \(\chi^2\)-statistic from part (c), and a critical value of \(3.84\), what is your conclusion?</p>
[[input:ans2d]] [[validation:ans2d]] [[feedback:prt_ans2d]]

<p>(e) Using your values from part (b), give a 95% confidence interval for the odds ratio, as a list \([\text{lower},\text{upper}]\).</p>
[[input:ans2e]] [[validation:ans2e]] [[feedback:prt_ans2e]]

<h4>Validity check</h4>

<p>(a) What is the smallest expected cell count in this table, to 2 decimal places?</p>
[[input:ans3a]] [[validation:ans3a]] [[feedback:prt_ans3a]]

<p>(b) Based on your answer to (a), is the chi-square approximation valid for this table?</p>
[[input:ans3b]] [[validation:ans3b]] [[feedback:prt_ans3b]]]]></text>
</questiontext>
<questionvariables>
<text><![CDATA[regime: rand([0,1]);

if regime=0 then (
  n_small_total: rand(80)+80,
  n_large_total: rand(80)+80,
  p_small_true: (rand(30)+20)/100,
  p_large_true: (rand(30)+5)/100
) else (
  n_small_total: rand(15)+10,
  n_large_total: rand(15)+10,
  p_small_true: (rand(40)+10)/100,
  p_large_true: (rand(40)+10)/100
);

B: round(n_small_total*p_small_true);
A: n_small_total - B;
D: round(n_large_total*p_large_true);
C: n_large_total - D;

s_assert(is(A>0 and B>0 and C>0 and D>0), true);

N: A+B+C+D;
n_s: A+B;
n_l: C+D;

p_hat_small: float(B/n_s);
p_hat_large: float(D/n_l);
diff_true: p_hat_small - p_hat_large;

p_pool: float((B+D)/N);
se_pooled: sqrt(p_pool*(1-p_pool)*(1/n_s+1/n_l));
z_true: float(diff_true/se_pooled);

se_unpooled: sqrt(p_hat_small*(1-p_hat_small)/n_s + p_hat_large*(1-p_hat_large)/n_l);
ci_diff_lower: float(diff_true - 1.96*se_unpooled);
ci_diff_upper: float(diff_true + 1.96*se_unpooled);

decision1_true: if abs(z_true)>1.96 then "reject" else "fail to reject";

odds_small_true: float(B/A);
odds_large_true: float(D/C);
OR_true: float(odds_small_true/odds_large_true);

chisq_true: float(N*(A*D-B*C)^2/((A+B)*(C+D)*(A+C)*(B+D)));
decision2_true: if chisq_true>3.84 then "reject" else "fail to reject";

se_logOR: sqrt(1/A+1/B+1/C+1/D);
ci_OR_lower: float(OR_true*exp(-1.96*se_logOR));
ci_OR_upper: float(OR_true*exp(1.96*se_logOR));

E_AA: float((A+B)*(A+C)/N);
E_AB: float((A+B)*(B+D)/N);
E_CA: float((C+D)*(A+C)/N);
E_CD: float((C+D)*(B+D)/N);
min_expected_true: min(E_AA,E_AB,E_CA,E_CD);

validity_true: if min_expected_true>=5 then "valid" else "invalid";

ta_mcq1a: [
  [1, true, "\\( \\hat p_{\\text{small}} - \\hat p_{\\text{large}} \\), the difference in proportions of *non-adopters* between groups"],
  [2, false, "\\( \\hat p_{\\text{large}} - \\hat p_{\\text{small}} \\), the difference in proportions of non-adopters between groups"],
  [3, false, "\\( \\hat p_{\\text{small}} - \\hat p_{\\text{large}} \\), the difference in proportions of *adopters* between groups"],
  [4, false, "The difference in raw counts of non-adopters between the two groups"]
];

ta_mcq2a: [
  [1, true, "\\( OR = \\dfrac{\\text{odds of non-adoption, small firms}}{\\text{odds of non-adoption, large firms}} \\)"],
  [2, false, "\\( OR = \\dfrac{\\text{odds of adoption, small firms}}{\\text{odds of adoption, large firms}} \\)"],
  [3, false, "\\( OR = \\dfrac{\\text{odds of non-adoption, large firms}}{\\text{odds of non-adoption, small firms}} \\)"],
  [4, false, "\\( OR = \\) the difference in proportions of non-adopters between the two groups"]
];

decision1_opts: [
  [1, is(decision1_true="reject"), "Reject \\(H_0\\): there is evidence of a difference in non-adoption proportions"],
  [2, is(decision1_true="fail to reject"), "Fail to reject \\(H_0\\): insufficient evidence of a difference in non-adoption proportions"]
];

decision2_opts: [
  [1, is(decision2_true="reject"), "Reject \\(H_0\\): there is evidence the odds ratio differs from 1"],
  [2, is(decision2_true="fail to reject"), "Fail to reject \\(H_0\\): insufficient evidence the odds ratio differs from 1"]
];

validity_opts: [
  [1, is(validity_true="valid"), "Valid — all expected cell counts are at least 5"],
  [2, is(validity_true="invalid"), "Invalid — at least one expected cell count is below 5"]
];]]></text>
</questionvariables>
<generalfeedback format="html">
<text><![CDATA[<p>The observed table was:</p>
<table border="1" cellpadding="6" style="border-collapse:collapse; text-align:center;">
<tr><th></th><th>Adopters</th><th>Non-adopters</th><th>Total</th></tr>
<tr><th>Small firms</th><td>{@A@}</td><td>{@B@}</td><td>{@A+B@}</td></tr>
<tr><th>Large firms</th><td>{@C@}</td><td>{@D@}</td><td>{@C+D@}</td></tr>
</table>

<h4>Thread 1: proportion difference</h4>
<p>\(\hat p_{\text{small}} = {@p_hat_small@}\), \(\hat p_{\text{large}} = {@p_hat_large@}\), so the difference is \({@diff_true@}\).</p>
<p>The pooled standard error (used for the hypothesis test) uses the pooled proportion \(\hat p = {@p_pool@}\), giving \(z = {@z_true@}\).</p>
<p>The (unpooled) standard error, used for the confidence interval, uses each group's own \(\hat p\) separately, giving a 95% CI of \([{@ci_diff_lower@}, {@ci_diff_upper@}]\).</p>
<p>Since \(|z| {@if abs(z_true)>1.96 then ">" else "\\le"@} 1.96\), we {@decision1_true@} \(H_0\).</p>

<h4>Thread 2: odds ratio</h4>
<p>\(\text{odds}_{\text{small}} = {@odds_small_true@}\), \(\text{odds}_{\text{large}} = {@odds_large_true@}\), so \(OR = {@OR_true@}\).</p>
<p>The chi-square statistic for this table is \(\chi^2 = {@chisq_true@}\), and the 95% CI for the odds ratio is \([{@ci_OR_lower@}, {@ci_OR_upper@}]\).</p>
<p>Since \(\chi^2 {@if chisq_true>3.84 then ">" else "\\le"@} 3.84\), we {@decision2_true@} \(H_0\).</p>

<p><strong>Key point:</strong> the pooled two-proportion \(z\)-test and the chi-square test of independence on a 2×2 table are mathematically equivalent — \(\chi^2 = z^2\) exactly, and \(1.96^2 = 3.8416 \approx 3.84\). So the two hypothesis-test decisions above will always agree. (Note the difference between the pooled standard error used for the test statistics and the unpooled standard error used for the confidence intervals — this is why a CI-based conclusion can occasionally look slightly different from the test-statistic conclusion near a boundary, even though the two hypothesis tests themselves always agree with each other.)</p>

<h4>Validity</h4>
<p>The smallest expected cell count is \({@min_expected_true@}\), so the chi-square approximation is {@validity_true@} here.</p>]]></text>
</generalfeedback>
<questionnote format="moodle_auto_format">
<text>\(A={@A@}, B={@B@}, C={@C@}, D={@D@}\), regime={@regime@}</text>
</questionnote>
<input>
<name>ans1a</name>
<type>radio</type>
<tans>ta_mcq1a</tans>
<showvalidation>0</showvalidation>
</input>
<input>
<name>ans1b</name>
<type>algebraic</type>
<tans>[p_hat_small, p_hat_large]</tans>
<forbidfloat>0</forbidfloat>
</input>
<input>
<name>ans1c</name>
<type>algebraic</type>
<tans>z_true</tans>
<forbidfloat>0</forbidfloat>
</input>
<input>
<name>ans1d</name>
<type>radio</type>
<tans>decision1_opts</tans>
<showvalidation>0</showvalidation>
</input>
<input>
<name>ans1e</name>
<type>algebraic</type>
<tans>[ci_diff_lower, ci_diff_upper]</tans>
<forbidfloat>0</forbidfloat>
</input>
<input>
<name>ans2a</name>
<type>radio</type>
<tans>ta_mcq2a</tans>
<showvalidation>0</showvalidation>
</input>
<input>
<name>ans2b</name>
<type>algebraic</type>
<tans>[odds_small_true, odds_large_true]</tans>
<forbidfloat>0</forbidfloat>
</input>
<input>
<name>ans2c</name>
<type>algebraic</type>
<tans>chisq_true</tans>
<forbidfloat>0</forbidfloat>
</input>
<input>
<name>ans2d</name>
<type>radio</type>
<tans>decision2_opts</tans>
<showvalidation>0</showvalidation>
</input>
<input>
<name>ans2e</name>
<type>algebraic</type>
<tans>[ci_OR_lower, ci_OR_upper]</tans>
<forbidfloat>0</forbidfloat>
</input>
<input>
<name>ans3a</name>
<type>algebraic</type>
<tans>min_expected_true</tans>
<forbidfloat>0</forbidfloat>
</input>
<input>
<name>ans3b</name>
<type>radio</type>
<tans>validity_opts</tans>
<showvalidation>0</showvalidation>
</input>
<prt>
  <name>prt_ans1a</name>
  <node>
    <name>0</name>
    <answertest>AlgEquiv</answertest>
    <sans>true</sans>
    <tans>true</tans>
    <trueanswernote>prt_ans1a-1-T</trueanswernote>
    <falseanswernote>prt_ans1a-1-F</falseanswernote>
  </node>
</prt>

<prt>
  <name>prt_ans1b</name>
  <node>
    <name>0</name>
    <answertest>AlgEquiv</answertest>
    <sans>true</sans>
    <tans>true</tans>
    <trueanswernote>prt_ans1b-1-T</trueanswernote>
    <falseanswernote>prt_ans1b-1-F</falseanswernote>
  </node>
</prt>

<prt>
  <name>prt_ans1c</name>
  <node>
    <name>0</name>
    <answertest>AlgEquiv</answertest>
    <sans>true</sans>
    <tans>true</tans>
    <trueanswernote>prt_ans1c-1-T</trueanswernote>
    <falseanswernote>prt_ans1c-1-F</falseanswernote>
  </node>
</prt>

<prt>
  <name>prt_ans1d</name>
  <node>
    <name>0</name>
    <answertest>AlgEquiv</answertest>
    <sans>true</sans>
    <tans>true</tans>
    <trueanswernote>prt_ans1d-1-T</trueanswernote>
    <falseanswernote>prt_ans1d-1-F</falseanswernote>
  </node>
</prt>

<prt>
  <name>prt_ans1e</name>
  <node>
    <name>0</name>
    <answertest>AlgEquiv</answertest>
    <sans>true</sans>
    <tans>true</tans>
    <trueanswernote>prt_ans1e-1-T</trueanswernote>
    <falseanswernote>prt_ans1e-1-F</falseanswernote>
  </node>
</prt>

<prt>
  <name>prt_ans2a</name>
  <node>
    <name>0</name>
    <answertest>AlgEquiv</answertest>
    <sans>true</sans>
    <tans>true</tans>
    <trueanswernote>prt_ans2a-1-T</trueanswernote>
    <falseanswernote>prt_ans2a-1-F</falseanswernote>
  </node>
</prt>

<prt>
  <name>prt_ans2b</name>
  <node>
    <name>0</name>
    <answertest>AlgEquiv</answertest>
    <sans>true</sans>
    <tans>true</tans>
    <trueanswernote>prt_ans2b-1-T</trueanswernote>
    <falseanswernote>prt_ans2b-1-F</falseanswernote>
  </node>
</prt>

<prt>
  <name>prt_ans2c</name>
  <node>
    <name>0</name>
    <answertest>AlgEquiv</answertest>
    <sans>true</sans>
    <tans>true</tans>
    <trueanswernote>prt_ans2c-1-T</trueanswernote>
    <falseanswernote>prt_ans2c-1-F</falseanswernote>
  </node>
</prt>

<prt>
  <name>prt_ans2d</name>
  <node>
    <name>0</name>
    <answertest>AlgEquiv</answertest>
    <sans>true</sans>
    <tans>true</tans>
    <trueanswernote>prt_ans2d-1-T</trueanswernote>
    <falseanswernote>prt_ans2d-1-F</falseanswernote>
  </node>
</prt>

<prt>
  <name>prt_ans2e</name>
  <node>
    <name>0</name>
    <answertest>AlgEquiv</answertest>
    <sans>true</sans>
    <tans>true</tans>
    <trueanswernote>prt_ans2e-1-T</trueanswernote>
    <falseanswernote>prt_ans2e-1-F</falseanswernote>
  </node>
</prt>

<prt>
  <name>prt_ans3a</name>
  <node>
    <name>0</name>
    <answertest>AlgEquiv</answertest>
    <sans>true</sans>
    <tans>true</tans>
    <trueanswernote>prt_ans3a-1-T</trueanswernote>
    <falseanswernote>prt_ans3a-1-F</falseanswernote>
  </node>
</prt>

<prt>
  <name>prt_ans3b</name>
  <node>
    <name>0</name>
    <answertest>AlgEquiv</answertest>
    <sans>true</sans>
    <tans>true</tans>
    <trueanswernote>prt_ans3b-1-T</trueanswernote>
    <falseanswernote>prt_ans3b-1-F</falseanswernote>
  </node>
</prt>
<qtest>
  <testcase>1</testcase>
  <testinput>
    <name>ans1a</name>
    <value>ta_mcq1a</value>
  </testinput>
  <testinput>
    <name>ans1b</name>
    <value>[p_hat_small, p_hat_large]</value>
  </testinput>
  <testinput>
    <name>ans1c</name>
    <value>z_true</value>
  </testinput>
  <testinput>
    <name>ans1d</name>
    <value>decision1_opts</value>
  </testinput>
  <testinput>
    <name>ans1e</name>
    <value>[ci_diff_lower, ci_diff_upper]</value>
  </testinput>
  <testinput>
    <name>ans2a</name>
    <value>ta_mcq2a</value>
  </testinput>
  <testinput>
    <name>ans2b</name>
    <value>[odds_small_true, odds_large_true]</value>
  </testinput>
  <testinput>
    <name>ans2c</name>
    <value>chisq_true</value>
  </testinput>
  <testinput>
    <name>ans2d</name>
    <value>decision2_opts</value>
  </testinput>
  <testinput>
    <name>ans2e</name>
    <value>[ci_OR_lower, ci_OR_upper]</value>
  </testinput>
  <testinput>
    <name>ans3a</name>
    <value>min_expected_true</value>
  </testinput>
  <testinput>
    <name>ans3b</name>
    <value>validity_opts</value>
  </testinput>
  <expected>
    <name>prt_ans1a</name>
    <expectedscore>1.0000000</expectedscore>
    <expectedpenalty>0.0000000</expectedpenalty>
    <expectedanswernote>prt_ans1a-1-T</expectedanswernote>
  </expected>
  <expected>
    <name>prt_ans1b</name>
    <expectedscore>1.0000000</expectedscore>
    <expectedpenalty>0.0000000</expectedpenalty>
    <expectedanswernote>prt_ans1b-1-T</expectedanswernote>
  </expected>
  <expected>
    <name>prt_ans1c</name>
    <expectedscore>1.0000000</expectedscore>
    <expectedpenalty>0.0000000</expectedpenalty>
    <expectedanswernote>prt_ans1c-1-T</expectedanswernote>
  </expected>
  <expected>
    <name>prt_ans1d</name>
    <expectedscore>1.0000000</expectedscore>
    <expectedpenalty>0.0000000</expectedpenalty>
    <expectedanswernote>prt_ans1d-1-T</expectedanswernote>
  </expected>
  <expected>
    <name>prt_ans1e</name>
    <expectedscore>1.0000000</expectedscore>
    <expectedpenalty>0.0000000</expectedpenalty>
    <expectedanswernote>prt_ans1e-1-T</expectedanswernote>
  </expected>
  <expected>
    <name>prt_ans2a</name>
    <expectedscore>1.0000000</expectedscore>
    <expectedpenalty>0.0000000</expectedpenalty>
    <expectedanswernote>prt_ans2a-1-T</expectedanswernote>
  </expected>
  <expected>
    <name>prt_ans2b</name>
    <expectedscore>1.0000000</expectedscore>
    <expectedpenalty>0.0000000</expectedpenalty>
    <expectedanswernote>prt_ans2b-1-T</expectedanswernote>
  </expected>
  <expected>
    <name>prt_ans2c</name>
    <expectedscore>1.0000000</expectedscore>
    <expectedpenalty>0.0000000</expectedpenalty>
    <expectedanswernote>prt_ans2c-1-T</expectedanswernote>
  </expected>
  <expected>
    <name>prt_ans2d</name>
    <expectedscore>1.0000000</expectedscore>
    <expectedpenalty>0.0000000</expectedpenalty>
    <expectedanswernote>prt_ans2d-1-T</expectedanswernote>
  </expected>
  <expected>
    <name>prt_ans2e</name>
    <expectedscore>1.0000000</expectedscore>
    <expectedpenalty>0.0000000</expectedpenalty>
    <expectedanswernote>prt_ans2e-1-T</expectedanswernote>
  </expected>
  <expected>
    <name>prt_ans3a</name>
    <expectedscore>1.0000000</expectedscore>
    <expectedpenalty>0.0000000</expectedpenalty>
    <expectedanswernote>prt_ans3a-1-T</expectedanswernote>
  </expected>
  <expected>
    <name>prt_ans3b</name>
    <expectedscore>1.0000000</expectedscore>
    <expectedpenalty>0.0000000</expectedpenalty>
    <expectedanswernote>prt_ans3b-1-T</expectedanswernote>
  </expected>
</qtest>
</question>
</quiz>