<?xml version="1.0" encoding="UTF-8"?>
<quiz>
<question type="stack">
<name>
<text>Corolla price</text>
</name>
<questiontext format="html">
<text><![CDATA[<p>A study of used Toyota Corollas relates the car's age (in years) to its advertised price (in $1000s), based on a sample of \(n=38\) cars. A least-squares regression of price on age gives the following software output:</p>
<p>
Intercept: \(b_0 = {@b0@}\)<br/>
Slope: \(b_1 = {@b1@}\)<br/>
SE(\(b_1\)) \(= {@se_b1@}\)<br/>
\(r = {@r@}\)
</p>
<p>We want to test \(H_0: \beta_1 = 0\) versus \(H_1: \beta_1 < 0\) (older cars are cheaper), using \(\alpha = 0.05\).</p>

<p><b>Part 1.</b> Compute the value of the test statistic \(t\) for this hypothesis test.</p>
[[input:ans_t]] [[validation:ans_t]] [[feedback:prt_ans_t]]

<p><b>Part 2.</b> Using \(df = 36\), compute the exact one-tailed p-value corresponding to your value of \(t\), and hence state the decision at \(\alpha = 0.05\).</p>
<p>p-value: [[input:ans_p]] [[validation:ans_p]] [[feedback:prt_ans_p]]</p>
<p>Decision: [[input:ans_decision]] [[validation:ans_decision]] [[feedback:prt_ans_decision]]</p>

<p><b>Part 3.</b> Compute \(R^2\), the percentage of variation in price explained by age, from the given value of \(r\). Give your answer as a percentage, to 1 decimal place.</p>
[[input:ans_r2]] [[validation:ans_r2]] [[feedback:prt_ans_r2]]

<p><b>Part 4.</b> Using the rule-of-thumb multiplier of 2, compute a 95% confidence interval for the population slope \(\beta_1\), using \(b_1 \pm 2\times SE(b_1)\).</p>
<p>Lower bound: [[input:ans_ci_lower]] [[validation:ans_ci_lower]] [[feedback:prt_ans_ci_lower]]</p>
<p>Upper bound: [[input:ans_ci_upper]] [[validation:ans_ci_upper]] [[feedback:prt_ans_ci_upper]]</p>]]></text>
</questiontext>
<questionvariables>
<text><![CDATA[load(distrib);
n: 38;
df: 36;
b1: float(-(rand(10)+5)/10);
tcase: rand(3);
t_target: if tcase=0 then 10+rand(5) elseif tcase=1 then 4+rand(4) else 2+rand(2);
se_b1: float(abs(b1)/t_target);
b0: float(rand(11)+12);
t_true: float(b1/se_b1);
r: float(-t_true/sqrt(df+t_true^2));
r2pct: float(100*r^2);
p_true: float(cdf_student_t(t_true, df));
ci_lower: float(b1 - 2*se_b1);
ci_upper: float(b1 + 2*se_b1);
ta_decision: [[reject, p_true<0.05, "Reject H0"],[noreject, p_true>=0.05, "Do not reject H0"]];]]></text>
</questionvariables>
<generalfeedback format="html">
<text><![CDATA[<p>The test statistic is
\[ t = \frac{b_1}{SE(b_1)} = \frac{ {@b1@} }{ {@se_b1@} } = {@t_true@}. \]</p>
<p>Since \(H_1: \beta_1 < 0\), the p-value is the one-tailed (lower-tail) probability
\[ p = P(T_{36} \le t) = {@p_true@}. \]
This is the ONE-tailed p-value, not the two-tailed value — a common mistake is to double it or use the wrong tail.</p>
[[if test="p_true<0.05"]]
<p>Since \(p < 0.05\), we reject \(H_0\): there is significant evidence that price decreases with age.</p>
[[else]]
<p>Since \(p \geq 0.05\), we do not reject \(H_0\): there is insufficient evidence that price decreases with age.</p>
[[/if]]
<p>\(R^2 = r^2 = {@r@}^2 = {@r2pct@}\%\) of the variation in price is explained by age. (Remember to report \(R^2\) as a percentage, not a proportion.)</p>
<p>The 95% confidence interval for \(\beta_1\) using the rule-of-thumb multiplier of 2 is
\[ b_1 \pm 2\times SE(b_1) = {@b1@} \pm 2({@se_b1@}) = ({@ci_lower@}, {@ci_upper@}). \]
Be careful to apply this margin to \(b_1\) only, not to \(b_0\) or \(r\).</p>]]></text>
</generalfeedback>
<questionnote format="moodle_auto_format">
<text>\(b_0={@b0@}, b_1={@b1@}, SE(b_1)={@se_b1@}, r={@r@}, t={@t_true@}, p={@p_true@}\)</text>
</questionnote>
<input>
<name>ans_t</name>
<type>numerical</type>
<tans>t_true</tans>
<forbidfloat>0</forbidfloat>
</input>
<input>
<name>ans_p</name>
<type>numerical</type>
<tans>p_true</tans>
<forbidfloat>0</forbidfloat>
</input>
<input>
<name>ans_decision</name>
<type>dropdown</type>
<tans>ta_decision</tans>
<showvalidation>0</showvalidation>
</input>
<input>
<name>ans_r2</name>
<type>numerical</type>
<tans>r2pct</tans>
<forbidfloat>0</forbidfloat>
</input>
<input>
<name>ans_ci_lower</name>
<type>numerical</type>
<tans>ci_lower</tans>
<forbidfloat>0</forbidfloat>
</input>
<input>
<name>ans_ci_upper</name>
<type>numerical</type>
<tans>ci_upper</tans>
<forbidfloat>0</forbidfloat>
</input>
<prt>
  <name>prt_ans_t</name>
  <node>
    <name>0</name>
    <answertest>AlgEquiv</answertest>
    <sans>true</sans>
    <tans>true</tans>
    <trueanswernote>prt_ans_t-1-T</trueanswernote>
    <falseanswernote>prt_ans_t-1-F</falseanswernote>
  </node>
</prt>

<prt>
  <name>prt_ans_p</name>
  <node>
    <name>0</name>
    <answertest>AlgEquiv</answertest>
    <sans>true</sans>
    <tans>true</tans>
    <trueanswernote>prt_ans_p-1-T</trueanswernote>
    <falseanswernote>prt_ans_p-1-F</falseanswernote>
  </node>
</prt>

<prt>
  <name>prt_ans_decision</name>
  <node>
    <name>0</name>
    <answertest>AlgEquiv</answertest>
    <sans>true</sans>
    <tans>true</tans>
    <trueanswernote>prt_ans_decision-1-T</trueanswernote>
    <falseanswernote>prt_ans_decision-1-F</falseanswernote>
  </node>
</prt>

<prt>
  <name>prt_ans_r2</name>
  <node>
    <name>0</name>
    <answertest>AlgEquiv</answertest>
    <sans>true</sans>
    <tans>true</tans>
    <trueanswernote>prt_ans_r2-1-T</trueanswernote>
    <falseanswernote>prt_ans_r2-1-F</falseanswernote>
  </node>
</prt>

<prt>
  <name>prt_ans_ci_lower</name>
  <node>
    <name>0</name>
    <answertest>AlgEquiv</answertest>
    <sans>true</sans>
    <tans>true</tans>
    <trueanswernote>prt_ans_ci_lower-1-T</trueanswernote>
    <falseanswernote>prt_ans_ci_lower-1-F</falseanswernote>
  </node>
</prt>

<prt>
  <name>prt_ans_ci_upper</name>
  <node>
    <name>0</name>
    <answertest>AlgEquiv</answertest>
    <sans>true</sans>
    <tans>true</tans>
    <trueanswernote>prt_ans_ci_upper-1-T</trueanswernote>
    <falseanswernote>prt_ans_ci_upper-1-F</falseanswernote>
  </node>
</prt>
<qtest>
  <testcase>1</testcase>
  <testinput>
    <name>ans_t</name>
    <value>t_true</value>
  </testinput>
  <testinput>
    <name>ans_p</name>
    <value>p_true</value>
  </testinput>
  <testinput>
    <name>ans_decision</name>
    <value>ta_decision</value>
  </testinput>
  <testinput>
    <name>ans_r2</name>
    <value>r2pct</value>
  </testinput>
  <testinput>
    <name>ans_ci_lower</name>
    <value>ci_lower</value>
  </testinput>
  <testinput>
    <name>ans_ci_upper</name>
    <value>ci_upper</value>
  </testinput>
  <expected>
    <name>prt_ans_t</name>
    <expectedscore>1.0000000</expectedscore>
    <expectedpenalty>0.0000000</expectedpenalty>
    <expectedanswernote>prt_ans_t-1-T</expectedanswernote>
  </expected>
  <expected>
    <name>prt_ans_p</name>
    <expectedscore>1.0000000</expectedscore>
    <expectedpenalty>0.0000000</expectedpenalty>
    <expectedanswernote>prt_ans_p-1-T</expectedanswernote>
  </expected>
  <expected>
    <name>prt_ans_decision</name>
    <expectedscore>1.0000000</expectedscore>
    <expectedpenalty>0.0000000</expectedpenalty>
    <expectedanswernote>prt_ans_decision-1-T</expectedanswernote>
  </expected>
  <expected>
    <name>prt_ans_r2</name>
    <expectedscore>1.0000000</expectedscore>
    <expectedpenalty>0.0000000</expectedpenalty>
    <expectedanswernote>prt_ans_r2-1-T</expectedanswernote>
  </expected>
  <expected>
    <name>prt_ans_ci_lower</name>
    <expectedscore>1.0000000</expectedscore>
    <expectedpenalty>0.0000000</expectedpenalty>
    <expectedanswernote>prt_ans_ci_lower-1-T</expectedanswernote>
  </expected>
  <expected>
    <name>prt_ans_ci_upper</name>
    <expectedscore>1.0000000</expectedscore>
    <expectedpenalty>0.0000000</expectedpenalty>
    <expectedanswernote>prt_ans_ci_upper-1-T</expectedanswernote>
  </expected>
</qtest>
</question>
</quiz>