<?xml version="1.0" encoding="UTF-8"?>
<quiz>
<question type="stack">
<name>
<text>Sample size mean1</text>
</name>
<questiontext format="html">
<text><![CDATA[<p>A sample is to be drawn from a population with standard deviation \(s = {@s@}\) kg, in order to estimate the population mean with a margin of error \(E\), using a 95% confidence interval (\(z = 1.96\)).</p>

<p><b>Part 1.</b> What is the minimum sample size \(n\) required to achieve a margin of error of \(E_1 = {@E1@}\) kg? (Give a whole number.)</p>
[[input:ans1]] [[validation:ans1]] [[feedback:prt_ans1]]

<p><b>Part 2.</b> Suppose instead the margin of error required is \(E_1/{@r2@} = {@E1/r2@}\) kg. Using your answer to Part 1, what is the new minimum sample size?</p>
[[input:ans2]] [[validation:ans2]] [[feedback:prt_ans2]]

<p><b>Part 3.</b> Suppose instead the margin of error required is \(E_1/{@r3@} = {@E1/r3@}\) kg. Using your answer to Part 1, what is the new minimum sample size?</p>
[[input:ans3]] [[validation:ans3]] [[feedback:prt_ans3]]

<p><b>Part 4.</b> For the same margin of error \(E_1 = {@E1@}\) kg, would a <b>smaller</b> or <b>larger</b> sample be needed if a {@conf_alt@}% confidence interval were used instead of a 95% confidence interval?</p>
[[input:ans4]] [[validation:ans4]] [[feedback:prt_ans4]]]]></text>
</questiontext>
<questionvariables>
<text>s: (5+rand(16))/10.0;
E1: rand([0.2,0.3,0.4,0.5]);
r2: rand([2,3,4]);
r3set: setdifference({2,3,4},{r2});
r3: rand(listify(r3set));
z95: 1.96;
n1: ceiling((z95*s/E1)^2);
n2: n1*r2^2;
n3: n1*r3^2;
dir: rand([1,2]);
conf_alt: if dir=1 then 90 else 99;
z_alt: if dir=1 then 1.645 else 2.576;
option_larger: larger;
option_smaller: smaller;
option_same: same_size;
option_cant: cant_tell;
mcq_ans: if dir=1 then [[option_larger,false,"Larger"],[option_smaller,true,"Smaller"],[option_same,false,"Same"],[option_cant,false,"Can't tell"]] else [[option_larger,true,"Larger"],[option_smaller,false,"Smaller"],[option_same,false,"Same"],[option_cant,false,"Can't tell"]];
mcq_ans: random_permutation(mcq_ans);
mcq_correct_ans: first(mcq_correct(mcq_ans));</text>
</questionvariables>
<generalfeedback format="html">
<text><![CDATA[<p><b>Part 1.</b> The required sample size is
\[ n = \left(\frac{z \cdot s}{E_1}\right)^2 = \left(\frac{1.96 \times {@s@}}{{@E1@}}\right)^2 = {@float((z95*s/E1)^2)@}, \]
which rounds up (since sample size must be a whole number) to \(n_1 = {@n1@}\).</p>

<p><b>Part 2.</b> Halving/scaling the margin of error to \(E_1/{@r2@}\) scales \(n\) by a factor of \({@r2@}^2 = {@r2^2@}\), since \(n\) is inversely proportional to \(E^2\). So the new required sample size is \(n_1 \times {@r2@}^2 = {@n1@} \times {@r2^2@} = {@n2@}\).</p>

<p><b>Part 3.</b> Similarly, scaling the margin of error to \(E_1/{@r3@}\) gives a new required sample size of \(n_1 \times {@r3@}^2 = {@n1@} \times {@r3^2@} = {@n3@}\).</p>

<p><b>Part 4.</b> A higher confidence level requires a larger critical value \(z\) (e.g. \(z = {@z_alt@}\) for {@conf_alt@}% confidence, compared to \(z=1.96\) for 95%). Since \(n\) is proportional to \(z^2\), a larger \(z\) requires a larger sample for the same margin of error, and a smaller \(z\) requires a smaller sample.</p>]]></text>
</generalfeedback>
<questionnote format="moodle_auto_format">
<text>\(s={@s@}, E_1={@E1@}, r_2={@r2@}, r_3={@r3@}, n_1={@n1@}, \text{dir}={@dir@}\)</text>
</questionnote>
<input>
<name>ans1</name>
<type>numerical</type>
<tans>n1</tans>
<forbidfloat>1</forbidfloat>
</input>
<input>
<name>ans2</name>
<type>numerical</type>
<tans>n2</tans>
<forbidfloat>1</forbidfloat>
</input>
<input>
<name>ans3</name>
<type>numerical</type>
<tans>n3</tans>
<forbidfloat>1</forbidfloat>
</input>
<input>
<name>ans4</name>
<type>radio</type>
<tans>mcq_ans</tans>
<showvalidation>0</showvalidation>
</input>
<prt>
  <name>prt_ans1</name>
  <node>
    <name>0</name>
    <answertest>AlgEquiv</answertest>
    <sans>ans1</sans>
    <tans>n1</tans>
    <trueanswernote>prt_ans1-1-T</trueanswernote>
    <falseanswernote>prt_ans1-1-F</falseanswernote>
  </node>
</prt>

<prt>
  <name>prt_ans2</name>
  <node>
    <name>0</name>
    <answertest>AlgEquiv</answertest>
    <sans>ans2</sans>
    <tans>ans1*r2^2</tans>
    <trueanswernote>prt_ans2-1-T</trueanswernote>
    <falsenextnode>1</falsenextnode>
    <falseanswernote>prt_ans2-1-F</falseanswernote>
  </node>
  <node>
    <name>1</name>
    <answertest>AlgEquiv</answertest>
    <sans>ans2</sans>
    <tans>ans1*r2</tans>
    <truescore>0</truescore>
    <trueanswernote>prt_ans2-2-T</trueanswernote>
    <truefeedback format="html">
      <text><![CDATA[<p>It looks like you scaled the sample size by \(r_2\) instead of \(r_2^2\) — since \(n\) depends on \(E^2\), scaling the margin of error by a factor scales \(n\) by the square of that factor. (This is checked against your own Part 1 answer.)</p>]]></text>
    </truefeedback>
    <falseanswernote>prt_ans2-2-F</falseanswernote>
    <falsefeedback format="html">
      <text><![CDATA[<p>This does not follow from your Part 1 answer using the correct scaling relationship \(n = n_1 \times r_2^2\). Check the scaling and re-apply it to your own value of \(n_1\).</p>]]></text>
    </falsefeedback>
  </node>
</prt>

<prt>
  <name>prt_ans3</name>
  <node>
    <name>0</name>
    <answertest>AlgEquiv</answertest>
    <sans>ans3</sans>
    <tans>ans1*r3^2</tans>
    <trueanswernote>prt_ans3-1-T</trueanswernote>
    <falsenextnode>1</falsenextnode>
    <falseanswernote>prt_ans3-1-F</falseanswernote>
  </node>
  <node>
    <name>1</name>
    <answertest>AlgEquiv</answertest>
    <sans>ans3</sans>
    <tans>ans1*r3</tans>
    <truescore>0</truescore>
    <trueanswernote>prt_ans3-2-T</trueanswernote>
    <truefeedback format="html">
      <text><![CDATA[<p>It looks like you scaled the sample size by \(r_3\) instead of \(r_3^2\) — since \(n\) depends on \(E^2\), scaling the margin of error by a factor scales \(n\) by the square of that factor. (This is checked against your own Part 1 answer.)</p>]]></text>
    </truefeedback>
    <falseanswernote>prt_ans3-2-F</falseanswernote>
    <falsefeedback format="html">
      <text><![CDATA[<p>This does not follow from your Part 1 answer using the correct scaling relationship \(n = n_1 \times r_3^2\). Check the scaling and re-apply it to your own value of \(n_1\).</p>]]></text>
    </falsefeedback>
  </node>
</prt>

<prt>
  <name>prt_ans4</name>
  <node>
    <name>0</name>
    <answertest>AlgEquiv</answertest>
    <sans>ans4</sans>
    <tans>mcq_correct_ans</tans>
    <trueanswernote>prt_ans4-1-T</trueanswernote>
    <falseanswernote>prt_ans4-1-F</falseanswernote>
  </node>
</prt>
<qtest>
  <testcase>1</testcase>
  <description>All parts correct</description>
  <testinput>
    <name>ans1</name>
    <value>n1</value>
  </testinput>
  <testinput>
    <name>ans2</name>
    <value>n1*r2^2</value>
  </testinput>
  <testinput>
    <name>ans3</name>
    <value>n1*r3^2</value>
  </testinput>
  <testinput>
    <name>ans4</name>
    <value>mcq_correct_ans</value>
  </testinput>
  <expected>
    <name>prt_ans1</name>
    <expectedscore>1.0000000</expectedscore>
    <expectedpenalty>0.0000000</expectedpenalty>
    <expectedanswernote>prt_ans1-1-T</expectedanswernote>
  </expected>
  <expected>
    <name>prt_ans2</name>
    <expectedscore>1.0000000</expectedscore>
    <expectedpenalty>0.0000000</expectedpenalty>
    <expectedanswernote>prt_ans2-1-T</expectedanswernote>
  </expected>
  <expected>
    <name>prt_ans3</name>
    <expectedscore>1.0000000</expectedscore>
    <expectedpenalty>0.0000000</expectedpenalty>
    <expectedanswernote>prt_ans3-1-T</expectedanswernote>
  </expected>
  <expected>
    <name>prt_ans4</name>
    <expectedscore>1.0000000</expectedscore>
    <expectedpenalty>0.0000000</expectedpenalty>
    <expectedanswernote>prt_ans4-1-T</expectedanswernote>
  </expected>
</qtest>

<qtest>
  <testcase>2</testcase>
  <description>Wrong part 1, but parts 2 and 3 correctly follow through from the student's own (wrong) part-1 value</description>
  <testinput>
    <name>ans1</name>
    <value>n1+2</value>
  </testinput>
  <testinput>
    <name>ans2</name>
    <value>(n1+2)*r2^2</value>
  </testinput>
  <testinput>
    <name>ans3</name>
    <value>(n1+2)*r3^2</value>
  </testinput>
  <testinput>
    <name>ans4</name>
    <value>mcq_correct_ans</value>
  </testinput>
  <expected>
    <name>prt_ans1</name>
    <expectedscore>0.0000000</expectedscore>
    <expectedpenalty>0.1000000</expectedpenalty>
    <expectedanswernote>prt_ans1-1-F</expectedanswernote>
  </expected>
  <expected>
    <name>prt_ans2</name>
    <expectedscore>1.0000000</expectedscore>
    <expectedpenalty>0.0000000</expectedpenalty>
    <expectedanswernote>prt_ans2-1-T</expectedanswernote>
  </expected>
  <expected>
    <name>prt_ans3</name>
    <expectedscore>1.0000000</expectedscore>
    <expectedpenalty>0.0000000</expectedpenalty>
    <expectedanswernote>prt_ans3-1-T</expectedanswernote>
  </expected>
  <expected>
    <name>prt_ans4</name>
    <expectedscore>1.0000000</expectedscore>
    <expectedpenalty>0.0000000</expectedpenalty>
    <expectedanswernote>prt_ans4-1-T</expectedanswernote>
  </expected>
</qtest>

<qtest>
  <testcase>3</testcase>
  <description>Correct part 1, but part 2 uses the linear-scaling mistake (r2 instead of r2^2)</description>
  <testinput>
    <name>ans1</name>
    <value>n1</value>
  </testinput>
  <testinput>
    <name>ans2</name>
    <value>n1*r2</value>
  </testinput>
  <testinput>
    <name>ans3</name>
    <value>n1*r3^2</value>
  </testinput>
  <testinput>
    <name>ans4</name>
    <value>mcq_correct_ans</value>
  </testinput>
  <expected>
    <name>prt_ans1</name>
    <expectedscore>1.0000000</expectedscore>
    <expectedpenalty>0.0000000</expectedpenalty>
    <expectedanswernote>prt_ans1-1-T</expectedanswernote>
  </expected>
  <expected>
    <name>prt_ans2</name>
    <expectedscore>0.0000000</expectedscore>
    <expectedpenalty>0.1000000</expectedpenalty>
    <expectedanswernote>prt_ans2-2-T</expectedanswernote>
  </expected>
  <expected>
    <name>prt_ans3</name>
    <expectedscore>1.0000000</expectedscore>
    <expectedpenalty>0.0000000</expectedpenalty>
    <expectedanswernote>prt_ans3-1-T</expectedanswernote>
  </expected>
  <expected>
    <name>prt_ans4</name>
    <expectedscore>1.0000000</expectedscore>
    <expectedpenalty>0.0000000</expectedpenalty>
    <expectedanswernote>prt_ans4-1-T</expectedanswernote>
  </expected>
</qtest>

<qtest>
  <testcase>4</testcase>
  <description>Correct part 1, but part 2 is an unrelated wrong answer (not the follow-through, not the linear-scaling mistake)</description>
  <testinput>
    <name>ans1</name>
    <value>n1</value>
  </testinput>
  <testinput>
    <name>ans2</name>
    <value>n1*r2^2+7</value>
  </testinput>
  <testinput>
    <name>ans3</name>
    <value>n1*r3^2</value>
  </testinput>
  <testinput>
    <name>ans4</name>
    <value>mcq_correct_ans</value>
  </testinput>
  <expected>
    <name>prt_ans1</name>
    <expectedscore>1.0000000</expectedscore>
    <expectedpenalty>0.0000000</expectedpenalty>
    <expectedanswernote>prt_ans1-1-T</expectedanswernote>
  </expected>
  <expected>
    <name>prt_ans2</name>
    <expectedscore>0.0000000</expectedscore>
    <expectedpenalty>0.1000000</expectedpenalty>
    <expectedanswernote>prt_ans2-2-F</expectedanswernote>
  </expected>
  <expected>
    <name>prt_ans3</name>
    <expectedscore>1.0000000</expectedscore>
    <expectedpenalty>0.0000000</expectedpenalty>
    <expectedanswernote>prt_ans3-1-T</expectedanswernote>
  </expected>
  <expected>
    <name>prt_ans4</name>
    <expectedscore>1.0000000</expectedscore>
    <expectedpenalty>0.0000000</expectedpenalty>
    <expectedanswernote>prt_ans4-1-T</expectedanswernote>
  </expected>
</qtest>

<qtest>
  <testcase>5</testcase>
  <description>Wrong MCQ answer for part 4 (opposite option to the correct one)</description>
  <testinput>
    <name>ans1</name>
    <value>n1</value>
  </testinput>
  <testinput>
    <name>ans2</name>
    <value>n1*r2^2</value>
  </testinput>
  <testinput>
    <name>ans3</name>
    <value>n1*r3^2</value>
  </testinput>
  <testinput>
    <name>ans4</name>
    <value>if mcq_correct_ans=option_larger then option_smaller else option_larger</value>
  </testinput>
  <expected>
    <name>prt_ans1</name>
    <expectedscore>1.0000000</expectedscore>
    <expectedpenalty>0.0000000</expectedpenalty>
    <expectedanswernote>prt_ans1-1-T</expectedanswernote>
  </expected>
  <expected>
    <name>prt_ans2</name>
    <expectedscore>1.0000000</expectedscore>
    <expectedpenalty>0.0000000</expectedpenalty>
    <expectedanswernote>prt_ans2-1-T</expectedanswernote>
  </expected>
  <expected>
    <name>prt_ans3</name>
    <expectedscore>1.0000000</expectedscore>
    <expectedpenalty>0.0000000</expectedpenalty>
    <expectedanswernote>prt_ans3-1-T</expectedanswernote>
  </expected>
  <expected>
    <name>prt_ans4</name>
    <expectedscore>0.0000000</expectedscore>
    <expectedpenalty>0.1000000</expectedpenalty>
    <expectedanswernote>prt_ans4-1-F</expectedanswernote>
  </expected>
</qtest>
</question>
</quiz>