<?xml version="1.0" encoding="UTF-8"?>
<quiz>
<question type="stack">
<name>
<text>One mean ci bears</text>
</name>
<questiontext format="html">
<text><![CDATA[<p>A biologist is studying the weight of adult male black bears. A random sample of \(n={@n@}\) male bears has sample mean weight \(\bar{x}={@xbar@}\) kg and sample standard deviation \(s={@s@}\) kg. The biologist wants to construct a 95% confidence interval for the population mean weight \(\mu\) of all adult male black bears, using the critical value \(z^*=1.96\).</p>

<p>Before doing any calculations, think about what population parameter \(\mu\) represents here, and picture the sampling distribution of \(\bar{x}\) that the confidence interval is based on. (This part is not graded.)</p>

<p><b>Part 1.</b> Calculate the standard error of the mean, \(SE=\dfrac{s}{\sqrt{n}}\). Give your answer to 3 decimal places.</p>
[[input:ans1]] [[validation:ans1]] [[feedback:prt_ans1]]

<p><b>Part 2.</b> Calculate the margin of error, \(ME=z^*\times SE\), using your own value of \(SE\) from Part 1. Give your answer to 3 decimal places.</p>
[[input:ans2]] [[validation:ans2]] [[feedback:prt_ans2]]

<p><b>Part 3.</b> Using your own value of \(ME\) from Part 2, calculate the lower and upper bounds of the confidence interval, \(\bar{x}-ME\) and \(\bar{x}+ME\). Give each answer to 2 decimal places.</p>
<p>Lower bound: [[input:ans3]] [[validation:ans3]] [[feedback:prt_ans3]]</p>
<p>Upper bound: [[input:ans4]] [[validation:ans4]] [[feedback:prt_ans4]]</p>

<p><b>Part 4.</b> Is it valid to use the large-sample \(z\)-confidence interval procedure here (i.e. is \(n\ge 25\))?</p>
[[input:ans5]] [[validation:ans5]] [[feedback:prt_ans5]]

<p><b>Part 5.</b> Which of the following is the correct interpretation of this confidence interval?</p>
[[input:ans6]] [[validation:ans6]] [[feedback:prt_ans6]]]]></text>
</questiontext>
<questionvariables>
<text><![CDATA[n: rand([10,15,18,22,30,45,80,150,250]);
xbar: rand(51)+60;
s: rand(31)+30;
z: 1.96;
se: s/sqrt(n);
me: z*se;
lower: xbar-me;
upper: xbar+me;
se_val: float(se);
me_val: float(me);
lower_val: float(lower);
upper_val: float(upper);
valid_p: is(n>=25);
mcq_valid: [[true, valid_p, "Yes, the procedure is valid"], [false, not valid_p, "No, the procedure is not valid"]];
mcq_interp: [
  [1, false, "There is a 95% probability that the true population mean \\(\\mu\\) lies between the lower and upper bound computed from this particular sample."],
  [2, true, "We are 95% confident that the interval from the lower to the upper bound contains the true population mean weight \\(\\mu\\) of all adult male black bears."],
  [3, false, "95% of all adult male black bears have weights between the lower and upper bound."],
  [4, false, "95% of all possible sample means \\(\\bar{x}\\) from repeated samples will fall between the lower and upper bound computed from this particular sample."]
];]]></text>
</questionvariables>
<generalfeedback format="html">
<text><![CDATA[<p>With \(n={@n@}\), \(\bar{x}={@xbar@}\), \(s={@s@}\):</p>
<p>\(SE=\dfrac{s}{\sqrt{n}}={@s@}/\sqrt{@n@}={@se_val@}\)</p>
<p>\(ME=z^*\times SE=1.96\times {@se_val@}={@me_val@}\)</p>
<p>Confidence interval: \(\bar{x}\pm ME = ({@lower_val@},\,{@upper_val@})\)</p>
[[if test="valid_p"]]
<p>Since \(n={@n@}\ge 25\), the sample size is large enough for the Central Limit Theorem to justify treating the sampling distribution of \(\bar{x}\) as approximately normal, so the \(z\)-interval procedure is valid here.</p>
[[else]]
<p>Since \(n={@n@}<25\), the sample size is too small to rely on the Central Limit Theorem to justify the normal approximation used in this \(z\)-interval procedure, so strictly this method should not be used here.</p>
[[/if]]
<p>The correct interpretation is: we are 95% confident that the interval contains the true population mean \(\mu\) — the confidence level describes the long-run success rate of the method, not a probability statement about this one fixed interval, and it says nothing directly about individual bears or individual sample means.</p>]]></text>
</generalfeedback>
<questionnote format="moodle_auto_format">
<text>\(n={@n@}, \bar{x}={@xbar@}, s={@s@}\), valid={@valid_p@}</text>
</questionnote>
<input>
<name>ans1</name>
<type>numerical</type>
<tans>se</tans>
<boxsize>10</boxsize>
<forbidfloat>0</forbidfloat>
</input>
<input>
<name>ans2</name>
<type>numerical</type>
<tans>me</tans>
<boxsize>10</boxsize>
<forbidfloat>0</forbidfloat>
</input>
<input>
<name>ans3</name>
<type>numerical</type>
<tans>lower</tans>
<boxsize>10</boxsize>
<forbidfloat>0</forbidfloat>
</input>
<input>
<name>ans4</name>
<type>numerical</type>
<tans>upper</tans>
<boxsize>10</boxsize>
<forbidfloat>0</forbidfloat>
</input>
<input>
<name>ans5</name>
<type>radio</type>
<tans>mcq_valid</tans>
<showvalidation>0</showvalidation>
</input>
<input>
<name>ans6</name>
<type>radio</type>
<tans>mcq_interp</tans>
<showvalidation>0</showvalidation>
</input>
<prt>
  <name>prt_ans1</name>
  <feedbackvariables>
    <text><![CDATA[tol: 0.0005;
correct_p: is(abs(ans1-se_val)<tol);
mistake_p: is(abs(ans1-float(s/n))<tol);]]></text>
  </feedbackvariables>
  <node>
    <name>0</name>
    <answertest>AlgEquiv</answertest>
    <sans>correct_p</sans>
    <tans>true</tans>
    <trueanswernote>prt_ans1-1-T</trueanswernote>
    <falsenextnode>1</falsenextnode>
    <falseanswernote>prt_ans1-1-F</falseanswernote>
  </node>
  <node>
    <name>1</name>
    <answertest>AlgEquiv</answertest>
    <sans>mistake_p</sans>
    <tans>true</tans>
    <truescore>0</truescore>
    <truepenalty>0.1</truepenalty>
    <trueanswernote>prt_ans1-2-T</trueanswernote>
    <truefeedback format="html">
      <text><![CDATA[<p>It looks like you divided by \(n\) instead of \(\sqrt{n}\) — remember \(SE=s/\sqrt{n}\), not \(s/n\).</p>]]></text>
    </truefeedback>
    <falseanswernote>prt_ans1-2-F</falseanswernote>
    <falsefeedback format="html">
      <text><![CDATA[<p>Not correct — check your calculation of \(SE=s/\sqrt{n}\).</p>]]></text>
    </falsefeedback>
  </node>
</prt>

<prt>
  <name>prt_ans2</name>
  <feedbackvariables>
    <text><![CDATA[tol2: 0.0005;
me_from_ans1: float(z*ans1);
correct_p: is(abs(ans2-me_from_ans1)<tol2);
mistake_p: is(abs(ans2-float(ans1))<tol2);]]></text>
  </feedbackvariables>
  <node>
    <name>0</name>
    <answertest>AlgEquiv</answertest>
    <sans>correct_p</sans>
    <tans>true</tans>
    <trueanswernote>prt_ans2-1-T</trueanswernote>
    <falsenextnode>1</falsenextnode>
    <falseanswernote>prt_ans2-1-F</falseanswernote>
  </node>
  <node>
    <name>1</name>
    <answertest>AlgEquiv</answertest>
    <sans>mistake_p</sans>
    <tans>true</tans>
    <truescore>0</truescore>
    <truepenalty>0.1</truepenalty>
    <trueanswernote>prt_ans2-2-T</trueanswernote>
    <truefeedback format="html">
      <text><![CDATA[<p>It looks like you used your \(SE\) directly as the margin of error — remember to multiply by \(z^*=1.96\).</p>]]></text>
    </truefeedback>
    <falseanswernote>prt_ans2-2-F</falseanswernote>
    <falsefeedback format="html">
      <text><![CDATA[<p>Not correct — check \(ME=z^*\times SE\), using your own \(SE\) from Part 1.</p>]]></text>
    </falsefeedback>
  </node>
</prt>

<prt>
  <name>prt_ans3</name>
  <feedbackvariables>
    <text><![CDATA[tol3: 0.005;
lower_from_ans2: float(xbar-ans2);
correct_p: is(abs(ans3-lower_from_ans2)<tol3);
signerror_p: is(abs(ans3-float(xbar+ans2))<tol3);]]></text>
  </feedbackvariables>
  <node>
    <name>0</name>
    <answertest>AlgEquiv</answertest>
    <sans>correct_p</sans>
    <tans>true</tans>
    <trueanswernote>prt_ans3-1-T</trueanswernote>
    <falsenextnode>1</falsenextnode>
    <falseanswernote>prt_ans3-1-F</falseanswernote>
  </node>
  <node>
    <name>1</name>
    <answertest>AlgEquiv</answertest>
    <sans>signerror_p</sans>
    <tans>true</tans>
    <truescore>0</truescore>
    <truepenalty>0.1</truepenalty>
    <trueanswernote>prt_ans3-2-T</trueanswernote>
    <truefeedback format="html">
      <text><![CDATA[<p>This matches the upper bound, not the lower bound — the lower bound is \(\bar{x}-ME\).</p>]]></text>
    </truefeedback>
    <falseanswernote>prt_ans3-2-F</falseanswernote>
    <falsefeedback format="html">
      <text><![CDATA[<p>Not correct — check your arithmetic for \(\bar{x}-ME\), using your own \(ME\) from Part 2.</p>]]></text>
    </falsefeedback>
  </node>
</prt>

<prt>
  <name>prt_ans4</name>
  <feedbackvariables>
    <text><![CDATA[tol4: 0.005;
upper_from_ans2: float(xbar+ans2);
correct_p: is(abs(ans4-upper_from_ans2)<tol4);
signerror_p: is(abs(ans4-float(xbar-ans2))<tol4);]]></text>
  </feedbackvariables>
  <node>
    <name>0</name>
    <answertest>AlgEquiv</answertest>
    <sans>correct_p</sans>
    <tans>true</tans>
    <trueanswernote>prt_ans4-1-T</trueanswernote>
    <falsenextnode>1</falsenextnode>
    <falseanswernote>prt_ans4-1-F</falseanswernote>
  </node>
  <node>
    <name>1</name>
    <answertest>AlgEquiv</answertest>
    <sans>signerror_p</sans>
    <tans>true</tans>
    <truescore>0</truescore>
    <truepenalty>0.1</truepenalty>
    <trueanswernote>prt_ans4-2-T</trueanswernote>
    <truefeedback format="html">
      <text><![CDATA[<p>This matches the lower bound, not the upper bound — the upper bound is \(\bar{x}+ME\).</p>]]></text>
    </truefeedback>
    <falseanswernote>prt_ans4-2-F</falseanswernote>
    <falsefeedback format="html">
      <text><![CDATA[<p>Not correct — check your arithmetic for \(\bar{x}+ME\), using your own \(ME\) from Part 2.</p>]]></text>
    </falsefeedback>
  </node>
</prt>

<prt>
  <name>prt_ans5</name>
  <node>
    <name>0</name>
    <answertest>CasEqual</answertest>
    <sans>ans5</sans>
    <tans>valid_p</tans>
    <trueanswernote>prt_ans5-1-T</trueanswernote>
    <falseanswernote>prt_ans5-1-F</falseanswernote>
  </node>
</prt>

<prt>
  <name>prt_ans6</name>
  <node>
    <name>0</name>
    <answertest>CasEqual</answertest>
    <sans>ans6</sans>
    <tans>2</tans>
    <trueanswernote>prt_ans6-1-T</trueanswernote>
    <falsenextnode>1</falsenextnode>
    <falseanswernote>prt_ans6-1-F</falseanswernote>
  </node>
  <node>
    <name>1</name>
    <answertest>CasEqual</answertest>
    <sans>ans6</sans>
    <tans>1</tans>
    <truescore>0</truescore>
    <truepenalty>0.1</truepenalty>
    <trueanswernote>prt_ans6-2-T</trueanswernote>
    <truefeedback format="html">
      <text><![CDATA[<p>This describes a probability statement about a fixed interval containing a fixed parameter — that's not what "95% confidence" means. The confidence level describes the long-run success rate of the method.</p>]]></text>
    </truefeedback>
    <falsenextnode>2</falsenextnode>
    <falseanswernote>prt_ans6-2-F</falseanswernote>
  </node>
  <node>
    <name>2</name>
    <answertest>CasEqual</answertest>
    <sans>ans6</sans>
    <tans>3</tans>
    <truescore>0</truescore>
    <truepenalty>0.1</truepenalty>
    <trueanswernote>prt_ans6-3-T</trueanswernote>
    <truefeedback format="html">
      <text><![CDATA[<p>This describes individual bears, but the interval is about the population mean weight \(\mu\), not individual weights.</p>]]></text>
    </truefeedback>
    <falsenextnode>3</falsenextnode>
    <falseanswernote>prt_ans6-3-F</falseanswernote>
  </node>
  <node>
    <name>3</name>
    <answertest>CasEqual</answertest>
    <sans>ans6</sans>
    <tans>4</tans>
    <truescore>0</truescore>
    <truepenalty>0.1</truepenalty>
    <trueanswernote>prt_ans6-4-T</trueanswernote>
    <truefeedback format="html">
      <text><![CDATA[<p>This describes the sampling distribution of \(\bar{x}\) across repeated samples, not what a single confidence interval tells us about \(\mu\).</p>]]></text>
    </truefeedback>
    <falseanswernote>prt_ans6-4-F</falseanswernote>
    <falsefeedback format="html">
      <text><![CDATA[<p>Not correct — review what "95% confidence" actually means for this interval.</p>]]></text>
    </falsefeedback>
  </node>
</prt>
<qtest>
  <testcase>1</testcase>
  <description>Fully correct answers throughout</description>
  <testinput>
    <name>ans1</name>
    <value>se_val</value>
  </testinput>
  <testinput>
    <name>ans2</name>
    <value>me_val</value>
  </testinput>
  <testinput>
    <name>ans3</name>
    <value>lower_val</value>
  </testinput>
  <testinput>
    <name>ans4</name>
    <value>upper_val</value>
  </testinput>
  <testinput>
    <name>ans5</name>
    <value>valid_p</value>
  </testinput>
  <testinput>
    <name>ans6</name>
    <value>2</value>
  </testinput>
  <expected>
    <name>prt_ans1</name>
    <expectedscore>1.0000000</expectedscore>
    <expectedpenalty>0.0000000</expectedpenalty>
    <expectedanswernote>prt_ans1-1-T</expectedanswernote>
  </expected>
  <expected>
    <name>prt_ans2</name>
    <expectedscore>1.0000000</expectedscore>
    <expectedpenalty>0.0000000</expectedpenalty>
    <expectedanswernote>prt_ans2-1-T</expectedanswernote>
  </expected>
  <expected>
    <name>prt_ans3</name>
    <expectedscore>1.0000000</expectedscore>
    <expectedpenalty>0.0000000</expectedpenalty>
    <expectedanswernote>prt_ans3-1-T</expectedanswernote>
  </expected>
  <expected>
    <name>prt_ans4</name>
    <expectedscore>1.0000000</expectedscore>
    <expectedpenalty>0.0000000</expectedpenalty>
    <expectedanswernote>prt_ans4-1-T</expectedanswernote>
  </expected>
  <expected>
    <name>prt_ans5</name>
    <expectedscore>1.0000000</expectedscore>
    <expectedpenalty>0.0000000</expectedpenalty>
    <expectedanswernote>prt_ans5-1-T</expectedanswernote>
  </expected>
  <expected>
    <name>prt_ans6</name>
    <expectedscore>1.0000000</expectedscore>
    <expectedpenalty>0.0000000</expectedpenalty>
    <expectedanswernote>prt_ans6-1-T</expectedanswernote>
  </expected>
</qtest>

<qtest>
  <testcase>2</testcase>
  <description>Classic SE mistake (s/n instead of s/sqrt(n))</description>
  <testinput>
    <name>ans1</name>
    <value>float(s/n)</value>
  </testinput>
  <expected>
    <name>prt_ans1</name>
    <expectedscore>0.0000000</expectedscore>
    <expectedpenalty>0.1000000</expectedpenalty>
    <expectedanswernote>prt_ans1-2-T</expectedanswernote>
  </expected>
</qtest>

<qtest>
  <testcase>3</testcase>
  <description>Follow-through — consistently wrong SE (forgot sqrt) should still earn ME/bounds credit</description>
  <testinput>
    <name>ans1</name>
    <value>float(s/n)</value>
  </testinput>
  <testinput>
    <name>ans2</name>
    <value>float(z*(s/n))</value>
  </testinput>
  <testinput>
    <name>ans3</name>
    <value>float(xbar-z*(s/n))</value>
  </testinput>
  <testinput>
    <name>ans4</name>
    <value>float(xbar+z*(s/n))</value>
  </testinput>
  <expected>
    <name>prt_ans1</name>
    <expectedscore>0.0000000</expectedscore>
    <expectedpenalty>0.1000000</expectedpenalty>
    <expectedanswernote>prt_ans1-2-T</expectedanswernote>
  </expected>
  <expected>
    <name>prt_ans2</name>
    <expectedscore>1.0000000</expectedscore>
    <expectedpenalty>0.0000000</expectedpenalty>
    <expectedanswernote>prt_ans2-1-T</expectedanswernote>
  </expected>
  <expected>
    <name>prt_ans3</name>
    <expectedscore>1.0000000</expectedscore>
    <expectedpenalty>0.0000000</expectedpenalty>
    <expectedanswernote>prt_ans3-1-T</expectedanswernote>
  </expected>
  <expected>
    <name>prt_ans4</name>
    <expectedscore>1.0000000</expectedscore>
    <expectedpenalty>0.0000000</expectedpenalty>
    <expectedanswernote>prt_ans4-1-T</expectedanswernote>
  </expected>
</qtest>

<qtest>
  <testcase>4</testcase>
  <description>Correct SE/ME but sign-flipped bounds — must be marked wrong independent of Parts 1-2</description>
  <testinput>
    <name>ans1</name>
    <value>se_val</value>
  </testinput>
  <testinput>
    <name>ans2</name>
    <value>me_val</value>
  </testinput>
  <testinput>
    <name>ans3</name>
    <value>float(xbar+me)</value>
  </testinput>
  <testinput>
    <name>ans4</name>
    <value>float(xbar-me)</value>
  </testinput>
  <expected>
    <name>prt_ans1</name>
    <expectedscore>1.0000000</expectedscore>
    <expectedpenalty>0.0000000</expectedpenalty>
    <expectedanswernote>prt_ans1-1-T</expectedanswernote>
  </expected>
  <expected>
    <name>prt_ans2</name>
    <expectedscore>1.0000000</expectedscore>
    <expectedpenalty>0.0000000</expectedpenalty>
    <expectedanswernote>prt_ans2-1-T</expectedanswernote>
  </expected>
  <expected>
    <name>prt_ans3</name>
    <expectedscore>0.0000000</expectedscore>
    <expectedpenalty>0.1000000</expectedpenalty>
    <expectedanswernote>prt_ans3-2-T</expectedanswernote>
  </expected>
  <expected>
    <name>prt_ans4</name>
    <expectedscore>0.0000000</expectedscore>
    <expectedpenalty>0.1000000</expectedpenalty>
    <expectedanswernote>prt_ans4-2-T</expectedanswernote>
  </expected>
</qtest>

<qtest>
  <testcase>5</testcase>
  <description>Validity check graded strictly off actual n, regardless of an unrelated wrong SE</description>
  <testinput>
    <name>ans1</name>
    <value>float(s/n)</value>
  </testinput>
  <testinput>
    <name>ans5</name>
    <value>valid_p</value>
  </testinput>
  <expected>
    <name>prt_ans1</name>
    <expectedscore>0.0000000</expectedscore>
    <expectedpenalty>0.1000000</expectedpenalty>
    <expectedanswernote>prt_ans1-2-T</expectedanswernote>
  </expected>
  <expected>
    <name>prt_ans5</name>
    <expectedscore>1.0000000</expectedscore>
    <expectedpenalty>0.0000000</expectedpenalty>
    <expectedanswernote>prt_ans5-1-T</expectedanswernote>
  </expected>
</qtest>

<qtest>
  <testcase>6</testcase>
  <description>Validity check answered incorrectly (opposite of actual n)</description>
  <testinput>
    <name>ans5</name>
    <value>not valid_p</value>
  </testinput>
  <expected>
    <name>prt_ans5</name>
    <expectedscore>0.0000000</expectedscore>
    <expectedpenalty>0.1000000</expectedpenalty>
    <expectedanswernote>prt_ans5-1-F</expectedanswernote>
  </expected>
</qtest>

<qtest>
  <testcase>7</testcase>
  <description>Interpretation — probability misinterpretation distractor</description>
  <testinput>
    <name>ans6</name>
    <value>1</value>
  </testinput>
  <expected>
    <name>prt_ans6</name>
    <expectedscore>0.0000000</expectedscore>
    <expectedpenalty>0.1000000</expectedpenalty>
    <expectedanswernote>prt_ans6-2-T</expectedanswernote>
  </expected>
</qtest>

<qtest>
  <testcase>8</testcase>
  <description>Interpretation — individual-bears misinterpretation distractor</description>
  <testinput>
    <name>ans6</name>
    <value>3</value>
  </testinput>
  <expected>
    <name>prt_ans6</name>
    <expectedscore>0.0000000</expectedscore>
    <expectedpenalty>0.1000000</expectedpenalty>
    <expectedanswernote>prt_ans6-3-T</expectedanswernote>
  </expected>
</qtest>

<qtest>
  <testcase>9</testcase>
  <description>Interpretation — repeated-sample-means misinterpretation distractor</description>
  <testinput>
    <name>ans6</name>
    <value>4</value>
  </testinput>
  <expected>
    <name>prt_ans6</name>
    <expectedscore>0.0000000</expectedscore>
    <expectedpenalty>0.1000000</expectedpenalty>
    <expectedanswernote>prt_ans6-4-T</expectedanswernote>
  </expected>
</qtest>
</question>
</quiz>